```python
import numpy as np
```

# 4. Conditional Independence in Bayesian Networks (6 points)

Consider the following Bayesian network:

<img src="bn.png" style="height:430px" />

<div class="alert alert-warning">
Indicate which of the following statements are true. $A \perp B$ means A is independent of B, $A \perp B \mid C$ means A is independent of B given C. (1 point / answer).
</div>

Store `False` into the corresponding result variable if the answer must be false and `True` otherwise.

*Hint: Read about **D-separation** and the **Bayes-ball algorithm** if you cannot intuitively solve these problems.* 

 1. $S \perp I$ ?
 2. $L \perp S$ ?
 3. $E \perp V$ ?
 4. $V \perp S \mid P, I$ ?
 5. $E \perp P \mid L$ ?
 6. $I \perp L \mid P, O$ ?

<!-- grade_id: cell-b2674b4fd61a5424, points: 0, grade: false, solution: true -->
```python
# Provide either True or False as an answer
results = {}
results['4.1'] = None
results['4.2'] = None
results['4.3'] = None
results['4.4'] = None
results['4.5'] = None
results['4.6'] = None

# YOUR CODE HERE
#raise NotImplementedError()

results['4.1'] = False
results['4.2'] = False
results['4.3'] = False
results['4.4'] = True
results['4.5'] = True
results['4.6'] = True
```

<!-- grade_id: cell-c3b644843fab77f4, points: 1, grade: true, solution: false -->
```python
assert results.get("4.1") is not None
```

<!-- grade_id: cell-13f728a3f45423a1, points: 1, grade: true, solution: false -->
```python
assert results.get("4.2") is not None
```

<!-- grade_id: cell-bb6b5105e346f780, points: 1, grade: true, solution: false -->
```python
assert results.get("4.3") is not None
```

<!-- grade_id: cell-e17fd6fbd55cab1c, points: 1, grade: true, solution: false -->
```python
assert results.get("4.4") is not None
```

<!-- grade_id: cell-26465c26053f9953, points: 1, grade: true, solution: false -->
```python
assert results.get("4.5") is not None
```

<!-- grade_id: cell-4e281ffec32cf722, points: 1, grade: true, solution: false -->
```python
assert results.get("4.6") is not None
```